Structure of the human CTF18–RFC clamp loader bound to PCNA

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public review):

Summary:

The authors report the structure of the human CTF18-RFC complex bound to PCNA. Similar structures (and more) have been reported by the O'Donnell and Li labs. This study should add to our understanding of CTF18-RFC in DNA replication and clamp loaders in general. However, there are numerous major issues that I recommend the authors fix.

Strengths:

The structures reported are strong and useful for comparison with other clamp loader structures that have been reported lately.

Comments on revisions:

The revised manuscript is greatly improved. The comparison with hRFC and the addition of direct PCNA loading data from the Hedglin group are particular highlights. I think this is a strong addition to the literature.

We thank the reviewer for their positive comments.

I only have minor comments on the revised manuscript.

(1) The clamp loading kinetic data in Figure 6 would be more easily interpreted if the three graphs all had the same x axes, and if addition of RFC was t=0 rather than t=60 sec.

We now analyze and plot EFRET as a function of time after complex addition, effectively setting the loader addition to t = 0 for each trace (Figure 6 and Figs S10-14 in the new manuscript). Baseline (Ymin) and plateau (Ymax) EFRET values were obtained by averaging the stable signal regions immediately before and after clamp-loader addition, respectively. Traces are normalized to their own dynamic range before fitting.

(2) The author's statement that "CTF18-RFC displayed a slightly faster rate than RFC" seems to me a bit misleading, even though this is technically correct. The two loaders have indistinguishable rate constants for the fast phase, and RFC is a bit slower than CTF18-RFC in the slow phase. However, the data also show that RFC is overall more efficient than CTF18-RFC at loading PCNA because much more flux through the fast phase (rel amplitudes 0.73 vs 0.36). Because the slow phase represents such a reduced fraction of loading events, the slight reduction in rate constant for the slow phase doesn't impact RFC's overall loading. And because the majority of loading events are in the fast phase, RFC has a faster halftime than CTF18-RFC. (Is it known what the different phases correspond to? If it is known, it might be interesting to discuss.)

We removed the quoted statement. We avoid comparing amplitude partitions (A₁/A_T) for CTF18-RFC because (i) a substantial fraction of the reaction occurs within the <7 s dead time, and (ii) single- vs double-exponential identifiability differs across complexes. Instead, we report model-minimal progress times: RFC t0.5 ≤ 7 s (faster onset), CTF18-RFC ~ 8 s, CTF18Δ165–194-RFC ~ 12 s; completion (t0.95): RFC ≈ 77 s, CTF18-RFC ≈ 77 s, mutant ≈ 145 s. This shows RFC has the steeper onset, while CTF18-RFC catches up in completion, and the mutant is slower overall. We briefly note that RFC’s phases have been assigned in prior stopped-flow work and are consistent with a rapid entry step and a slower repositioning/complex release phase; we do not assign phases for CTF18-RFC here and instead rely on model-minimal timing comparisons to avoid over-interpretation.

(3) AAA+ is an acronym for "ATPases Associated with diverse cellular Activities" rather than "Adenosine Triphosphatase Associated".

Corrected to ATPases Associated with diverse cellular Activities (AAA+).

Reviewer #2 (Public review):

Summary

Briola and co-authors have performed a structural analysis of the human CTF18 clamp loader bound to PCNA. The authors purified the complexes and formed a complex in solution. They used cryo-EM to determine the structure to high resolution. The complex assumed an auto-inhibited conformation, where DNA binding is blocked, which is of regulatory importance and suggests that additional factors could be required to support PCNA loading on DNA. The authors carefully analysed the structure and compared it to RFC and related structures.

Strength & Weakness

Their overall analysis is of high quality, and they identified, among other things, a humanspecific beta-hairpin in Ctf18 that flexible tethers Ctf18 to Rfc2-5. Indeed, deletion of the beta-hairpin resulted in reduced complex stability and a reduction in a primer extension assay with Pol ε. Moreover, the authors identify that the Ctf18 ATP-binding domain assumes a more flexible organisation.

The data are discussed accurately and relevantly, which provides an important framework for rationalising the results.

All in all, this is a high-quality manuscript that identifies a key intermediate in CTF18-dependent clamp loading.

Comments on revisions:

The authors have done a nice job with the revision.

We thank the reviewer for their very positive comments.

Reviewer #3 (Public review):

Summary:

CTF18-RFC is an alternative eukaryotic PCNA sliding clamp loader which is thought to specialize in loading PCNA on the leading strand. Eukaryotic clamp loaders (RFC complexes) have an interchangeable large subunit which is responsible for their specialized functions. The authors show that the CTF18 large subunit has several features responsible for its weaker PCNA loading activity, and that the resulting weakened stability of the complex is compensated by a novel beta hairpin backside hook. The authors show this hook is required for the optimal stability and activity of the complex.

Relevance:

The structural findings are important for understanding RFC enzymology and novel ways that the widespread class of AAA ATPases can be adapted to specialized functions. A better understanding of CTF18-RFC function will also provide clarity into aspects of DNA replication, cohesion establishment and the DNA damage response.

Strengths:

The cryo-EM structures are of high quality enabling accurate modelling of the complex and providing a strong basis for analyzing differences and similarities with other RFC complexes.

Weaknesses:

The manuscript would have benefited from a more detailed biochemical analysis using mutagenesis and assays to tease apart the differences with the canonical RFC complex. Analysis of the FRET assay could be improved.

Overall appraisal:

Overall, the work presented here is solid and important. The data is mostly sufficient to support the stated conclusions.

We thank the reviewer for their mainly positive assessment. Following this reviewer suggestion, we have re-analysed the FRET assay data and amended the manuscript accordingly.

Comments on revisions:

While the authors addressed my previous specific concerns, they have now added a new experiment which raises new concerns.

The FRET clamp loading experiments (Fig. 6) appear to be overfitted so that the fitted values are unlikely to be robust and it is difficult to know what they mean, and this is not explained in this manuscript. Specifically, the contribution of two exponentials is floated in each experiment. By eye, CTF18-RFC looks much slower than RFC1-RFC (as also shown previously in the literature) but the kinetic constants and text suggest it is faster. This is because the contribution of the fast exponential is substantially decreased, and the rate constants then compensate for this. There is a similar change in contribution of the slow and fast rates between WT CTF18 and the variant (where the data curves look the same) and this has been balanced out by a change in the rate constants, which is then interpreted as a defect. I doubt the data are strong enough to confidently fit all these co-dependent parameters, especially for CTF18, where a fast initial phase is not visible. I would recommend either removing this figure or doing a more careful and thorough analysis.

We appreciate the reviewer’s concern regarding potential overfitting of the kinetic data in Figure 6. To address this, we performed a model-minimal re-analysis designed specifically to avoid parameter covariance and over-interpretation (Figure 6 and Figs S11-14 in the new manuscript). Only data recorded after the instrument’s <7 s dead time were included in the fits, thereby excluding the partially obscured early region of the reaction. For each clamp loader complex, we selected the minimal kinetic model that produced residuals randomly distributed about zero. This approach yielded a single-exponential fit for CTF18-RFC, whereas RFC and CTF18Δ165–194-RFC required double-exponential fits; single-exponential models for the latter two complexes left structured residuals, clearly indicating the presence of an additional kinetic phase.

Rather than relying on co-dependent amplitude and rate parameters, we quantified the reactions by reporting progress times (t0.5, t0.90, t0.95), which provide a model-independent measure of reaction speed. This directly addresses the reviewer’s concern and allows a fair comparison of the relative kinetics among the complexes.

From this analysis, RFC exhibited the fastest onset (t0.5 ≤ 7 s; lower bound), while CTF18RFC and CTF18Δ165–194-RFC showed progressively slower half-times of approximately 8 s and 12 s, respectively. Completion times further emphasized these differences: both RFC and CTF18-RFC reached 95 % completion at ~77 s, whereas the mutant required ~145 s. Despite these kinetic distinctions, CTF18-RFC and its β-hairpin deletion mutant achieved similar EFRET plateaus, indicating that the mutation slows reaction progression but does not reduce the overall extent of PCNA loading.

Finally, we emphasize that our interpretation is deliberately conservative. We do not assign distinct kinetic phases to CTF18-RFC, as their molecular basis remains unresolved. RFC’s phases have been characterized in prior stopped-flow studies, but CTF18-RFC likely follows a distinct or simplified pathway. Our conclusions are thus limited to what the data unambiguously support: deletion of the Ctf18 β-hairpin decreases the rate—but not the extent—of PCNA loading, consistent with the reduced stimulation of Pol ε primer extension observed under single-turnover conditions.

Changes from the earlier version to the revised version

Abstract

Sliding clamps like PCNA are crucial processivity factors for replicative polymerases, requiring specific clamp loaders for loading onto DNA. The human alternative clamp loader CTF18-RFC interacts with the leading strand polymerase Pol ε and loads PCNA onto primer/template DNA using its RFC pentameric module. Here, we provide a structural characterization of the human CTF18-RFC complex and its interaction with PCNA. Our cryo-EM data support that the Ctf8 and Dcc1 subunits of CTF18-RFC, which form the regulatory module interacting with Pol ε, are flexibly tethered to the RFC module. A 2.9 Å cryo-EM structure shows the RFC module bound to PCNA in an auto-inhibited conformation similar to the canonical RFC loader, marking the initial step of the clamp-loading reaction. The unique RFC1 (Ctf18) large subunit of CTF18-RFC, which based on the cryo-EM map shows high relative flexibility, is anchored to PCNA through an atypical low-affinity PIP box in the AAA+ domain and engages the RFC5 subunit using a novel β-hairpin at the disordered N-terminus. We show that deletion of this β-hairpin impairs the CTF18-RFC–PCNA complex stability, slows down clamp loading, and decreases the rate of primer synthesis by Pol ε. Our research identifies distinctive structural characteristics of the human CTF18-RFC complex, providing insights into its role in PCNA loading and the stimulation of leading strand synthesis by Pol ε.

Introduction

Earlier version: High fidelity and efficiency are critical for chromosomal DNA duplication in eukaryotes. This precision is achieved through the coordinated action of replicative DNA polymerases, Pol δ and Pol ε, anchored to the sliding clamp, Proliferating Cell Nuclear Antigen (PCNA) (1). The ring-shaped PCNA homotrimer encircles DNA, enhancing the processivity of these polymerases in DNA replication and repair processes, including excision repair and homologous recombination (2), and plays an important role in sister chromatid cohesion establishment (3). The closed-ring topology of PCNA necessitates the transient opening of one interface to facilitate the incorporation and trapping of primer-template DNA. This critical function is executed by pentameric AAA+ (Adenosine Triphosphatase Associated) ATPases, collectively known as clamp loaders (4–6). Extensive structural and biochemical work in different organisms (7–12) established a common mechanism for the clamp-loading reaction, which proceeds through three main steps (13–18). Initially, the clamp loader engages with the intact clamp ring, adopting an autoinhibited configuration that precludes DNA binding (10). This stage sets the groundwork for the subsequent transition, where the clamp loader alters its conformation to an activated state. This change facilitates the opening of the clamp ring, creating a gap sufficiently large to permit the passage of duplex DNA and aligning the clamp loader for primer-template DNA interaction. The final phase of this process is initiated by the binding of DNA to the clamp loader, catalysing ATP hydrolysis. DNA binding prompts the closure of the clamp and hydrolysis of ATP induces the concurrent disassembly of the closed clamp loader from the sliding clamp-DNA complex, completing the cycle necessary for the engagement of the replicative polymerases to start DNA synthesis.

Revised version: High fidelity and efficiency are critical for chromosomal DNA duplication in eukaryotes. This precision is achieved through the coordinated action of replicative DNA polymerases, Pol δ and Pol ε, anchored to the sliding clamp, Proliferating Cell Nuclear Antigen (PCNA) (1). The ring-shaped PCNA homotrimer encircles DNA, enhancing the processivity of these polymerases in DNA replication and repair processes, including excision repair and homologous recombination (2), and plays an important role in sister chromatid cohesion establishment (3). The closed-ring topology of PCNA necessitates the transient opening of one interface to facilitate the incorporation and trapping of primer-template DNA. This critical function is executed by pentameric AAA+ (ATPases Associated with diverse cellular Activities) ATPases, collectively known as clamp loaders (4–6). Extensive structural and biochemical work in different organisms (7–12) established a common mechanism for the clamp-loading reaction, which proceeds through three main steps (13–18). Initially, the clamp loader engages with the intact clamp ring, adopting an autoinhibited configuration that precludes DNA binding (10). This stage sets the groundwork for the subsequent transition, where the clamp loader alters its conformation to an activated state. This change facilitates the opening of the clamp ring, creating a gap sufficiently large to permit the passage of duplex DNA and aligning the clamp loader for primer-template DNA interaction. The final phase of this process is initiated by the binding of DNA to the clamp loader, catalysing ATP hydrolysis. DNA binding prompts the closure of the clamp and hydrolysis of ATP induces the concurrent disassembly of the closed clamp loader from the sliding clamp-DNA complex, completing the cycle necessary for the engagement of the replicative polymerases to start DNA synthesis.

The canonical Replication Factor C (RFC), consisting of a large RFC1 subunit and four subunits (RFC2, RFC3, RFC4, and RFC5), acts as the primary clamp loader for PCNA (4), (19). However, eukaryotes also employ alternative loaders (20), including CTF18-RFC (6), (21–24), which likely use a conserved loading mechanism but are functionally specialized through specific protein interactions and context-dependent roles in DNA replication.

Genome-wide analyses revealed that CTF18-RFC predominantly facilitates PCNA loading on the leading strand, in contrast to the canonical RFC complex, which mainly loads PCNA on the lagging strand (25). Although CTF18-RFC is not essential for bulk DNA replication in yeast (26, 27), it plays a vital role in ensuring sister chromatid cohesion (21), (28, 29). This may be due to its selective loading of PCNA linked to EcoI acetyltransferase (3), which acetylates Smc3, preventing cohesin destabilization by WapI (30–32). Recruitment of CTF18-RFC to the leading strand is also important for activation of the replication checkpoint (22) (24) (33, 34).

The CTF18-RFC complex features two distinct modules with separate functions (23), (28, 35, 36): a catalytic RFC module for PCNA loading, consisting of a large RFC1-like subunit (hereby referred to as Ctf18), alongside the four RFC2, RFC3, RFC4 and RFC5 subunits utilized by the canonical RFC complex, and a regulatory module, made up of the Ctf8 and Dcc1 subunits. These latter two subunits attach to the Ctf18 C-terminus, forming the Ctf18-1-8 module. A long linker predicted to be flexible connects these two segments (34) (37). Structural and functional studies have demonstrated that the interaction between the Ctf18-1-8 module and the catalytic domain of Pol2 (Pol2CAT), the principal subunit of Pol ε, directs CTF18-RFC to replication forks and boosts its clamp loading efficiency (34) (37). This interaction situates the clamp loader close to the primer/template junction, facilitating PCNA loading (34). Pol2CAT is tethered to the remainder of Pol ε via an unstructured linker, making it an integral component of the CMGE (Cdc45-Mcm-GINS-Pol ε) core replisome complex (38). It has been proposed that Pol ε utilizes both the CMG helicase and PCNA as processivity factors to facilitate normal replication rates (39) (40). Tethering by the CMG complex might allow Pol ε to dissociate from the 3′ end of the leading strand but stay at the replication fork until leading strand synthesis restarts. Given the relatively weak interaction of Pol ε with PCNA (41), it is possible that multiple PCNA loading events on the leading strand via CTF18-RFC are necessary to achieve full Pol ε processivity.

Although the structure of the Ctf18-1-8 module in association with Pol2CAT has been elucidated (34) (37), the architecture of the CTF18-RFC module and its interaction with PCNA remain uncharacterized. To investigate the clamp loading mechanism employed by CTF18-RFC, we reconstituted the full human CTF18-RFC complex with PCNA and determined its structure using cryo-electron microscopy (cryo-EM). Our cryo-EM data supports the prediction that the Ctf18-1-8 and RFC modules are flexibly tethered. The analysis yielded high-resolution reconstructions of the entire CTF18-RFC module bound to a closed PCNA ring. Although the overall architecture mirrors the autoinhibited state observed in the canonical RFC loader in a complex with PCNA (10), distinctive features within CTF18-RFC were identified. These distinctions may offer insights into the specialized functional roles CTF18-RFC plays in DNA replication.

Results

Structure of the human CTF18-RFC–PCNA complex

We employed single-particle cryo-EM to determine the structure of human CTF18-RFCbound to PCNA. To this purpose, the full CTF18-RFC complex, including Ctf8 and Dcc1, was expressed in insect cells and purified to homogeneity (Fig S1). The purified complex includes all subunits as shown in the SDS-PAGE gel (Fig S1). With the aim of trapping CTF18-RFCin different steps of the clamp loading reaction, the complex was vitrified in the presence of ATP, PCNA, and primer/template DNA but omitting the Mg2+ cofactor to halt hydrolysis (42), and then imaged by cryo-EM (Fig S2). Surprisingly, image processing yielded only a major 3D class, including the CTF18-RFC pentamer bound to a closed PCNA ring (Fig 1a-b, Fig S2-3), with no evidence of DNA density. The Ctf18-1-8 module, which is connected to the RFC module by a 91 amino acid linker (Fig 1a), is not resolved in the map (Fig 1b), suggesting it is flexibly tethered to the RFC module. The map was refined to a global resolution of 2.9 Å (Fig S2-3), which allowed model building of the five CTF18-RFC subunits as well as the PCNA homotrimer (Fig 1c). The final model refines to a map-to-model FSC of 3.0 Å at 0.5 threshold, with good refinement statistics (S1 Table).

Cryo-EM structure of the human CTF18-RFC–PCNA complex in the presence of ATP.

a) Domain organization of the CTF18-RFC subunits forming the RFC pentamer. The dotted line above the Ctf18 subunit highlights the regions that are observed in the cryo-EM map. b) Two views of the cryo-EM map of the CTF18-RFC–PCNA complex, coloured by subunits. The less sharp definition of the map region corresponding to the AAA+ domain of the Ctf18 subunit supports a relative high mobility of this subunit. c) Molecular model of the CTF18-RFC–PCNA complex in ribbon representation. d) Structures of the human (left) and Saccharomyces cerevisiae (right) RFC–PCNA complex, which present the RFC pentamer in an auto-inhibited conformation analogous to the one captured in the CTF18-RFC–PCNA complex.

The five CTF18-RFC subunits are positioned on top of the closed PCNA ring, which is planar and undistorted (Fig 1c). A structural comparison shows that the CTF18-RFC pentamer was trapped in an autoinhibited conformation analogous to that reported for the canonical human (h) and yeast (sc) RFC (8), (10), pertaining to the initial step of the clamp loading reaction, right before clamp opening (Fig 1d). In such conformation, the AAA+ domains form an overtwisted and asymmetric spiral that is incompatible with DNA binding (8), (10). The AAA+, collar, and A’ domains of the Ctf18 subunit (Fig 2a) align with the homologous domains of hRFC1 (Fig 2b). As in the hRFC structure (10), the Ctf18 subunit interacts with RFC2, RFC3, and PCNA-I. Three main deviations from the hRFC–PCNA structure are observed (10): a different tilt of the PCNA ring (Fig 1c-d), significant mobility of the large Ctf18 subunit (Fig 1b and Fig S2) and, importantly, the presence of a β-hairpin engaged to the outer surface of the RFC5 subunit, which we mapped to the disordered N-terminus of the Ctf18 subunit (Fig 2a and 2c). These features are analyzed in detail below.

Structure of the Ctf18 subunit.

a) Sequence of the human Ctf18 subunit, highlighting the secondary structural elements as assigned in the cryo-EM model. The atypical PCNA-interacting motif (PIP-box) and conserved DEXX box are boxed. Residues that are not modelled in the cryo-EM structure are labelled as X in grey. Below the Ctf18 sequence to the left, the sequences of the N-terminus of human and other lower eukaryotes Ctf18, together with the secondary prediction by AlphaFold (45) are shown, highlighting the conservation of the β-hairpin fold. Below, the cryo-EM β-hairpin fold model is overlaid onto the corresponding motifs predicted by AlphaFold for Drosophila or Saccharomyces cerevisiae Ctf18. Below the Ctf18 sequence to the right, the alignment of the C-terminal region of human and Saccharomyces cerevisiae Ctf18, which interacts with the Ctf8/Dcc1 subunits, is shown. Asterisks denote conserved residues. b) Superposition of the human Ctf18 and RFC1 (PDB: 6VVO) subunit structures, with subdomains labelled. c) Structure of the Ctf18 subunit, in ribbon representation, within the CTF18-RFC–PCNA complex. The Ctf18 subdomains are coloured-coded as in panel (a). The other complex subunits are shown as surfaces.

Structure of the large subunit of CTF18-RFC

The Ctf18 subunit presents 26% identity with the homologous subunit RFC1 of hRFC and an overall structural conservation, with some deviations particularly in the A’ domain (Fig 2b). Compared to the other subunits, the AAA+ domain of Ctf18 shows increased flexibility, as evidenced by the local resolution map of the complex as well as 3D variability analysis (43) (44) (Fig S2d and Supplementary Video 1). Increased flexibility was not detected in the RFC1 subunit of human RFC bound to PCNA (10). This discrepancy could be explained by diminished inter-subunit interactions involving both the AAA+ and collar domains of Ctf18 compared to RFC1, as illustrated in Fig S4: the AAA+ domain of Ctf18 forms fewer stabilizing salt bridges and h-bonds with RFC2 than RFC1 does, and the Ctf18 collar domain establishes markedly weaker interactions with both RFC2 and RFC3 than the extensive network observed for RFC1. Conversely, the interaction between the A’ domain of CTF18-RFC and RFC3 is strong as observed for RFC. Overall, the total buried surface area between Ctf18 and interacting subunits is comparable to that of RFC1 (∼5500 Å² versus ∼5300 Å²). However, the buried surface area between Ctf18 and RFC2 is markedly smaller than that of RFC1 (∼1400 Å² and ∼2300 Å²). As explained below, the lost interactions mediated by the AAA+ and collar domains of Ctf18 are partially counterbalanced by those established with the N-terminal β-hairpin, which is absent in the canonical RFC. Additionally, the interaction between Ctf18 and PCNA is weaker compared to the interaction between RFC1 and PCNA, further contributing to these structural differences (see next section). The Ctf18 C-terminal region (residues 953-975) predicted to interact with the Ctf8 and Dcc1 subunits (Fig 2a) is invisible in the map.

The most striking feature is a small β-hairpin-containing fold mapped between residues 165-194 in the disordered N-terminus (Fig 2a and 2c) of the Ctf18 subunit, ∼80 Å apart from the Ctf18 AAA+ domain, to which it is connected by a disordered linker of 86 amino acids (residues 195-280). This N-terminal fold plugs into a groove at the outer surface of the collar domain of the RFC5 subunit defined by helices α1, α3, α4, and α5 (Fig 3a). The resolution of the map in this region is sufficient for confident model building of the interface (Fig 3b). The interface buries 984 Å2 and is characterized by both apolar and polar interactions (Fig 3a). A set of hydrophobic side chains protruding from Ctf18 β-hairpin (V178, V180, L190, L192) approaches a complementary patch of hydrophobic residues on RFC5 (L255 in α1; F287, V291 and F293 in α3; I336 and V337 in α5). The interface is further strengthened by several hydrogen bonds mostly mediated by main-chain atoms in the Ctf18 β-hairpin, and by a salt bridge between R193 of the Ctf18 β-hairpin and D292 of RFC5. It is worth highlighting that RFC2 separates the Ctf18 and RFC5 subunits. Despite this separation, Ctf18 projects the β-hairpin directly onto RFC5, bypassing any interactions with RFC2 (Fig 2c). This observation supports that the β-hairpin functions as a “latch” to stabilize the CTF18-RFC structure.

Structural details of the Ctf18–RFC5 and CTF18-RFC–PCNA interactions.

a) Model of the RFC5 collar domain and Ctf18 N-terminal β-hairpin in ribbon representation. The inset shows the interface in stick representation, with hydrogen bonds shown as green dotted lines. b) Map region around the Ctf18 N-terminal β-hairpin. c) Model of the PCNA homotrimer bound to the four CTF18 interacting regions. To the right, the structure of the yeast RFC–PCNA complex with an open clamp, highlighting the engagement of all five RFC subunits to PCNA, required for clamp opening. d) Model of the regions of the CTF18-RFC–PCNA interactions, with CTF18-RFC subunit residues shown as sticks and PCNA as surface or ribbon. Polar interactions are shown as black dotted lines. The amino acid sequence encompassing the interacting motifs in the various subunits are shown.

Similarly to human CTF18-RFC, long regions predicted to be disordered by AlphaFold (45) are present in the N-terminal sequence of Ctf18 of lower eukaryotes. The β-hairpin region of human Ctf18 does not show strict sequence conservation, but a similar fold in the N-terminus of lower eukaryotes Ctf18 is predicted (Fig 2a). Therefore, lower eukaryotes may use a similar β-hairpin motif to bind the corresponding subunit of the RFC-module complex (RFC5 in human, Rfc3 in S. cerevisiae), emphasizing its importance.

Interactions of CTF18-RFC with PCNA

CTF18-RFC engages two of the three PCNA protomers through four contact points involving the AAA+ domains of Ctf18, RFC2, RFC4, and RFC5 (Fig 1c and Fig 3c). The interactions of PCNA with the RFC2 and RFC5 subunits are analogous to those reported in the hRFC–PCNA complex (10). However, while Ctf18 engages PCNA-I via a PIP-box interaction similar to hRFC1 (10), its atypical PIP motif — lacking aromatic residues (Fig. 2a and Fig. 3c) — results in a smaller buried surface area (∼780 Å² vs. ∼1010 Å²). A methionine residue (M419) inserts into the so-called “Q-pocket,” which is typically occupied by a glutamine in canonical PIP-boxes, while V422 occupies the conserved hydrophobic pocket that usually accommodates aromatic residues. This low-affinity interaction does not stably anchor Ctf18 to PCNA-I and may contribute to the relative flexibility of the AAA+ domain of Ctf18 observed in the complex (Fig S2 and Supplementary Video 1). The PIP-box of RFC5 is also atypical but contains two hydrophobic residues (I114 and F115) that insert into the conserved pocket of PCNA-II. The interacting motifs of RFC2 and RFC4 differ from each other in sequence but both engage similar regions involving the interdomain connecting loop (IDCL) of PCNA-I and PCNA-II, respectively (Fig 3d).

Interestingly, while RFC4 interacts with PCNA in the CTF18-RFC complex, it does not in the hRFC–PCNA complex (10). The RFC4–PCNA interaction, which was also observed in the structure of S. cerevisiae RFC bound to open PCNA (11), brings PCNA closer to the CTF18-RFC module, thereby explaining the distinct tilt of the PCNA ring observed in our structure compared to the hRFC–PCNA complex (Fig. 1c and 1d). Given the structure of scRFC bound to open PCNA (11), where all five RFC subunits are bound to PCNA (Fig 3d), we predict that a conformational shift in CTF18–RFC, specifically bringing the fifth subunit (RFC3) to engage PCNA-III, would be necessary to disrupt the interface between PCNA-I and PCNA-III, thereby opening the clamp and aligning the CTF18-RFC subunits for primer-template DNA interaction. The reason for the absence of such an active CTF18-RFC conformation in our cryo-EM analysis remains uncertain. However, similar to observations with the hRFC–PCNA complex, where an open-clamp state was not reported (10), this could be attributed to the hydrophobic regions of the open clamp interacting with the water-air interface, potentially leading to protein denaturation during the vitrification process to prepare the cryo-EM grids.

Nucleotide binding to the AAA+ modules

Earlier version: The map is consistent with nucleotides in four of the five AAA+ domains of the CTF18-RFC subunits (Fig 4a). ATP is bound to subunits Ctf18, RFC2, and RFC5, while ADP appears to engage RFC3, consistent with the structure of the hRFC−PCNA complex reconstituted with the slowly hydrolysable nucleotide ATPγS (10). Density for the adenine ring of ATP in the Ctf18 subunit is poor, likely a consequence of the increased mobility of the AAA+ domain of this subunit (Fig 4a). Compared to the canonical RFC1, the Walker

Revised version: The map is consistent with nucleotides in four of the five AAA+ domains of the CTF18-RFC subunits (Fig 4a). ATP is bound to subunits Ctf18, RFC2, and RFC5, while ADP appears to engage RFC3, consistent with the structure of the hRFC–PCNA complex reconstituted with the slowly hydrolysable nucleotide ATPγS (10). Density for the adenine ring of ATP in the Ctf18 subunit is poor, likely a consequence of the increased mobility of the AAA+ domain of this subunit (Fig 4a). Compared to the canonical RFC1, the Walker A motif of Ctf18 contains a leucine residue (L378) instead of a valine, but this substitution does not appear to affect its interaction with the ATP nucleotide (Fig S5). Since RFC3 lacks critical catalytic residues and bears a substitution in the Walker A motif (10), the presence of ADP, rather than a product of ATP hydrolysis, probably results from co-purification of the complex where ADP contaminants are present. In the human RFC complex, it has been previously shown that an intact nucleotide binding site in RFC3 has an important structural role in the complex assembly (46, 47). If this is the case for CTF18-RFC, it remains to be established, but it is likely that ADP binding to this subunit may play a stabilizing role in all RFC-like loaders.

Analysis of the auto-inhibited conformation of the CTF18 pentamer and nucleotide coordination.

a) Ribbon model of the CTF18-RFC pentamer (PCNA hidden). Nucleotide molecules are shown as sticks, with the corresponding cryo-EM density. b) Surface representation of (top) the CTF18-RFC pentamer and (bottom) the yeast RFC pentamer bound to open PCNA and primer-template DNA. PCNA is hidden in both structures. The overtwisted CTF18-RFC pentamer cannot accommodate DNA within the inner chamber. c) Model of the ATP active sites in the CTF18-RFC–PCNA complex, showing disruption of the interactions mediated by the arginine fingers. Below, insets show the nucleotide-binding site of subunit RFC4 in either the human CTF18-RFC (left) or human RFC (right) models, highlighting the structural incompatibility with nucleotide binding in the CTF18-RFC structure. Polar interactions with ATPγS in human RFC4 are indicated as cyan dotted lines. d) Detail of the inner chamber region of the CTF18-RFC pentamer, showing the β-hairpin (E-plug) of RFC5 protruding in the region predicted to be occupied by DNA. The DNA (grey ribbon) is modeled as in the yeast RFC–DNA–PCNA complex. e) The structures of the RFC1 subunit from the yeast RFC– DNA–PCNA complex (11) and the Ctf18 subunit from the CTF18-RFC–PCNA complex are aligned to highlight conservation of the separation pin. Side chains of residues implicated in this motif in yeast RFC, along with the corresponding amino acids in CTF18-RFC, are shown.

Removed in the revised version: A motif of Ctf18 contains a leucine residue (L378) instead of a valine, but this substitution does not appear to affect its interaction with the ATP nucleotide (Fig S5). Since RFC3 lacks critical catalytic residues and bears a substitution in the Walker A motif (10), the presence of ADP, rather than a product of ATP hydrolysis, probably results from co-purification of the complex where ADP contaminants are present. In the human RFC complex, it has been previously shown that an intact nucleotide binding site in RFC3 has an important structural role in the complex assembly (46, 47). If this is the case for CTF18-RFC, it remains to be established, but it is likely that ADP binding to this subunit may play a stabilizing role in all RFC-like loaders.

Differently from the hRFC–PCNA complex, where RFC4 is bound to ATPγS, in CTF18-RFC, this subunit is not engaged with a nucleotide (Fig 4a). The residues in the AAA+ of RFC4 of CTF18-RFC coordinated to ATPγS in hRFC are displaced and incompatible with nucleotide binding (Fig 4c). In particular, the side chains of R45 and P46 in RFC4 occupy the space where ATP would normally bind, thereby sterically hindering nucleotide accommodation in the CTF18-RFC structure (Fig 4c). Interestingly, mutation of the conserved K84 residue in the ATP-binding motif of RFC4 in hRFC did not impair the ability of this subunit to assemble with other RFC subunits (46, 47). In agreement, our structure shows that the lack of ATP binding to RFC4 does not preclude the assembly of the CTF18-RFC pentamer in the autoinhibited conformation. However, considering the importance of ATP binding of RFC4 for the hRFC clamp loading function (48–51), it is expected that ATP binding to RFC4 in CTF18-RFC is required to switch to an enzymatically active conformation.

The asymmetric spiral of the CTF18-RFC AAA+ pentamer (Fig 4a-b) disrupts the active sites at the Ctf18/RFC2, RFC2/5, and RFC5/4 interfaces (Fig 4c). At these interfaces, the trans-acting arginine fingers are positioned too far away to interact with the γ-phosphate of ATP (e.g., ∼9 Å for R184 of RFC2, ∼10 Å for R168 of RFC5, and ∼7 Å for R193 of RFC4) (Fig 4c). Prior research has indicated that hydrolysis at the large subunit/RFC2 interface is not essential for clamp loading by various loaders (48–51), while the others are critical for the clamp-loading activity of eukaryotic RFCs.

In CTF18-RFC, the AAA+ spiral is overtwisted in a manner reminiscent of the autoinhibited conformations observed in hRFC and scRFC (8) (10) (Fig 1d), leading to steric hindrance that impedes DNA binding (Fig 4b). Specifically, the overtwisted RFC3 and RFC4 subunits occupy the space normally filled by DNA. In the autoinhibited structure of hRFC, the β-hairpin (residues 68-88) located at the base of RFC3 also referred to as “E-plug”, extends towards the anticipated location of DNA and interacts with the Rossmann fold of RFC1 (10). Conversely, in CTF18-RFC, while this same β-hairpin is oriented toward the interior of the DNA-binding channel (Fig 4d), it does not make direct contact with the Ctf18 subunit. Collectively, the E-plug appears a conserved feature in eukaryotic clamp loaders, underlying its importance in maintaining the pentamer in its autoinhibited conformation, by preventing DNA access through the central channel. We observed that human CTF18-RFC shares an additional structural feature with both human and yeast RFC: the presence of a separation pin in the large subunit (10, 11). Although the separation pin is not essential for clamp loading (11), its conservation across species and clamp loader types suggests an auxiliary or context-dependent function. Interestingly, residues W638 and F582—proposed to stabilize the separation pin and its interaction with DNA in yeast RFC—are substituted by F692 and H645 in human CTF18-RFC, respectively (Fig 4e), possibly indicating a mechanistically similar role in both systems.

To discount the possibility that the absence of the Mg2+ ions in the cryo-EM sample could have hindered the proper assembly of the CTF18-RFCpentamer for DNA engagement, we used cryo-EM to visualize CTF18-RFC with PCNA, primer-template DNA, ATP, and Mg2+ (Fig S6-7, S1 Table). Despite the inclusion of Mg2+, the resulting 3.2 Å reconstruction of the complex remained in a configuration analogous to that without Mg2+ (Fig 5a-b, Fig S6). In this reconstruction, distinct density for Mg2+ is noted in the active sites of RFC2 and RFC5, coordinating ATP (Fig 5b). Similar to the map without Mg2+, ADP coordination is present in RFC3, while RFC4 is not engaged with any nucleotide. These findings collectively imply that the presence of Mg2+ in the nucleotide binding sites does not influence the transition of CTF18-RFC from an autoinhibited to an active state. This result also confirms the remarkable stability of the autoinhibited conformation, even in hydrolysing conditions.

Cryo-EM structure of the human CTF18-RFC–PCNA complex in the presence of ATP and Mg2+.

a) Two views of the Cryo-EM map, coloured by subunits. b) Molecular model in ribbon representation. The inset shows a map and model of the ATP molecule bound to the RFC2 subunit, with a discernible Mg2+ ion.

The Ctf18 β-hairpin plays a role in PCNA loading and in stimulating DNA synthesis by Pol ε

Based on the structures reported herein, the β-hairpin at the disordered N-terminal domain of Ctf18 is a distinctive feature that may serve to stabilize the CTF18-RFC pentamer and maximize its clamp loading efficiency. To investigate the functional importance of the β-hairpin, we purified a CTF18-RFC mutant with a truncation at the disordered N-terminal domain of the Ctf18 subunit, specifically, the segment spanning residues 165-194 (CTF18Δ165-194-RFC). This mutation eliminates the entire Ctf18 β-hairpin and is expected to destabilize the CTF18-RFC pentamer. In agreement, when mutant CTF18-RFC was imaged by cryo-EM with PCNA in the same conditions used for wild type CTF18-RFC, only a small subset of particles of the complex was observed, while the majority of particles contained dissociated PCNA (Fig S8-9). We hypothesized that this destabilizing mutation in CTF18-RFC may reduce its efficiency in loading PCNA onto DNA.

Added in the revised version: We directly compared clamp loading by CTF18-RFC, CTF18Δ165–194-RFC and the canonical RFC complex, using a pre–steady-state ensemble FRET assay to monitor the loading of Cy5-labeled PCNA onto a Cy3-labeled primer–template DNA substrate pre-coated with RPA (Fig 6a, Figs S10-14). Prior to the addition of clamp loaders (Fig S11a), donor (I563) and acceptor (I665) fluorescence intensities remained low and stable over time (Fig S11b–d). Here, I563 and I665 correspond to the fluorescence intensities of the Cy3 donor (563 nm) and Cy5 acceptor (665 nm), respectively, and were used to calculate the apparent FRET efficiency (EFRET). The constant EFRET values recorded during the first 60 s represent the experimental baseline (YMin) corresponding to the absence of interactions between Cy5–PCNA and the 5′ddPCy3/T·RPA nucleoprotein complex (52, 53). Upon addition of a clamp loader, I665 increased concomitantly with a decrease in I563, after which both signals stabilized, indicating the onset and persistence of FRET (Fig S11b–d). These synchronized, anti-correlated changes in donor and acceptor intensities report the loading of Cy5–PCNA onto the DNA substrate. After reaching a plateau, EFRET values remained constant for at least 60 s, representing the amount of loaded and stably bound PCNA rings (52). For each condition, EFRET values within the final plateau region were fit to flat lines to determine the maximum EFRET (YMax) reached after clamp loader addition. Importantly, CTF18-RFC and CTF18Δ165–194-RFC displayed similar dynamic ranges of EFRET values (i.e., YMax – YMin, Fig S11c-d), indicating that the β-hairpin deletion does not affect the overall extent of PCNA loading. All EFRET values measured after addition of each clamp loader were normalized to their respective dynamic ranges, defined by the baseline (YMin) and maximum (YMax) EFRET values, and plotted as a function of time following complex addition.

Removed in the revised version: To test this, and to directly compare CTF18-RFC activity to the canonical clamp loader RFC, we used a pre-steady state ensemble FRET assay to monitor loading of Cy5-labeled PCNA onto a Cy3-labeled primer-template DNA substrate pre-coated with RPA (Fig 6 and Fig S10). Upon addition of RFC, CTF18-RFC, or the mutant CTF18 Δ165-194 -RFC complex, a time-resolved increase in FRET signal was observed, corresponding to productive PCNA loading. While both CTF18 and CTF18Δ165-194 -RFC were able to load PCNA onto DNA, the loading kinetics differed significantly: CTF18-RFC displayed a slightly faster rate than RFC, whereas CTF18Δ165-194-RFC exhibited a >2-fold reduction in the fast-phase rate constant (k₁), consistent with a functional defect (Table 1). Importantly, the FRET amplitude was comparable between the two CTF18-RFC isoforms, indicating that the mutation does not prevent PCNA loading per se but affects the rate at which loading occurs. These results suggest that the β-hairpin contributes to the efficiency—but not the capacity—of PCNA loading by CTF18-RFC.

Earlier version: Loading of PCNA onto P/T junctions by human clamp loader complexes.

Revised version: FRET assays to monitor loading of PCNA onto P/T junctions by human clamp loader complexes.

Earlier version: a) Schematic representation of the FRET pair and experiment to monitor loading of Cy5-PCNA onto 5′ddPCy3/T DNA substrates engaged by RPA. The front and back faces of PCNA are displayed in green and grey, respectively. When loaded onto a 5′ddPCy3/T DNA substrate, the Cy5 label on the back face of PCNA is oriented towards the Cy3 label near the blunt duplex end of the 5′ddPCy3/T DNA substrate, yielding a FRET signal. b -d) Data. Each trace is the mean of at least three independent traces with the S.E.M. shown in grey. The time trajectories of I563 (magenta) and I665 (cyan) are displayed in the top panels and the corresponding EFRET (mustard) is displayed in the bottom panels. The time at which a clamp loader complex is added is indicated by a red arrow. Changes in I563 and I665 are indicated in the top panel by magenta and cyan arrows, respectively. The I563, I665, and the corresponding EFRET values observed prior to the addition of clamp loader complexes represents the complete absence of interactions between 5¢ddPCy3/T·RPA complexes and Cy5-PCNA. For observations, these values are each fit to flat line that is extrapolated to the axis limits. For observation, the I563 and I665 values observed after the addition of clamp loader complexes are fit to double exponential declines (I563) and rises (I665), respectively. The EFRET values observed after the addition of a given clamp loader are fit to double exponential rises and the respective rate constants (k1 and k2) are reported in the corresponding figure and in Table 1. Data for loading of PCNA by RFC, CTF18-RFC, and CTF18 Δ165-194 -RFC are displayed in panels b, c, and d, respectively.

Revised version: a) Schematic representation of the FRET pair and experiment to monitor loading of Cy5-PCNA onto 5′ddPCy3/T DNA substrates engaged by RPA. The front and back faces of PCNA are displayed in green and grey, respectively. When loaded onto a 5′ddPCy3/T DNA substrate, the Cy5 label on the back face of PCNA is oriented towards the Cy3 label near the blunt duplex end of the 5′ddPCy3/T DNA substrate, yielding a FRET signal b-d) Data. Each normalized EFRET trace is the mean of at least three independent traces with the S.E.M. shown in grey. Each respective trace is fit to its corresponding minimal kinetic model and the times required for the trace to reach 50% (t0.50), 90%, (t0.90), and 95% (t0.95) of the maximal value (i.e., 1.0) are reported. Data for loading of PCNA by RFC, CTF18, and CTF18Δ165-194-RFC are displayed in panels b, c, and d, respectively.

Added in the revised version: Under our experimental conditions, the clamp loader complexes must bind ATP and free PCNA before initiating loading, which can give rise to multiple kinetic phases. Such multiphasic behaviour has been described previously for RFC in stopped-flow experiments and is also evident in our ensemble FRET setup (52–54). The kinetic steps represented by each phase have been well defined for RFC but not for CTF18-RFC. In our measurements, the instrumental dead time was <7 s. Within this interval, a substantial fraction of PCNA loading (ranging from ∼30% for CTF18-RFC to ∼70% for RFC; Fig 6b-d) occurs, indicating clear kinetic differences among the clamp loader complexes but precluding full resolution of all reaction phases for CTF18-RFC.

Added in the revised version: To enable direct comparison of loading rates, only normalized EFRET values recorded after the dead time were fitted to the minimal kinetic model that yielded randomly distributed residuals. For CTF18-RFC, the data were best described by a single-exponential rise (Fig S12), whereas RFC and CTF18Δ165–194-RFC required a double-exponential fit (Fig S13a–b), consistent with the presence of an additional kinetic phase (Fig S13c–d). From these minimal models, we calculated the times required to reach 50%, 90%, and 95% of the maximal normalized EFRET signal (t0.50, t0.90, t0.95), corresponding respectively to half completion and near completion of PCNA loading (Fig 6b-d and Fig S14). For RFC, t0.50 occurred within the dead time and is thus reported as a lower limit (Fig 6b).

Added in the revised version: Across the three complexes, the kinetics of PCNA loading differed significantly (Fig 6b-d). RFC reached half completion within the experimental dead time (t0.50 ≤ 7 s), consistent with rapid loading, whereas CTF18-RFC and CTF18Δ165–194-RFC exhibited slower half-times of approximately 8 s and 12 s, respectively. Completion times further underscored these differences: both RFC and CTF18-RFC reached 95 % completion at ∼77 s, while the CTF18Δ165–194-RFC mutant required ∼145 s. This indicates that the β-hairpin deletion significantly slows the reaction progression. Altogether, the results from these FRET studies (Fig 6, Fig S10–S14) indicate that deletion of the β-hairpin from CTF18 does not affect the overall extent of PCNA loading but rather slows the reaction progression.

Earlier version: We then tested whether the observed loading defect observed with CTF18Δ165-194-RFC results in diminished stimulation of primer synthesis by Pol ε. To test this hypothesis, we performed primer extension assays using Pol ε, PCNA, and either wild-type or mutant CTF18-RFC constructs to assess their ability to stimulate DNA synthesis by Pol ε. Our results indicate a significant reduction in primer extension products when the β-hairpin was deleted (Fig 7a). Measurement of band intensities revealed an approximately two-fold decrease in the amount of extended primer produced by the mutant compared to the wild-type (Fig 7b). To further support that the observed effect on primer extension is due to the mutation, we conducted a concentration titration comparing wild-type and mutant CTF18-RFC (Fig S11). The CTF18-RFC mutant displayed altered concentration dependencies compared to the wild-type, thereby validating the mutant phenotype and underscoring the structural importance of the β-hairpin feature. Notably, the primer extension assay was conducted with 2 nM plasmid DNA and excess CTF18-RFC (40 nM), PCNA (100 nM), and Pol ε (30 nM), establishing single-turnover conditions per DNA molecule. Under these conditions, the rate—not the extent—of PCNA loading becomes the key determinant of synthesis dynamics. Although Pol ε can initiate synthesis independently of PCNA, the clamp is required for full processivity. Therefore, the >2-fold slower loading rate observed for the CTF18-RFC mutant likely delays formation of highly processive complexes, resulting in a slower accumulation of extended products and explaining the observed ∼2-fold reduction in synthesis.

Revised version: We then tested whether the observed loading defect observed with CTF18Δ165-194 -RFC results in diminished stimulation of primer synthesis by Pol ε. To test this hypothesis, we performed primer extension assays using Pol ε, PCNA, and either wild-type or mutant CTF18-RFC constructs to assess their ability to stimulate DNA synthesis by Pol ε. Our results indicate a significant reduction in primer extension products when the β-hairpin was deleted (Fig 7a). Measurement of band intensities revealed an approximately two-fold decrease in the amount of extended primer produced by the mutant compared to the wild-type (Fig 7b). To further support that the observed effect on primer extension is due to the mutation, we conducted a concentration titration comparing wild-type and mutant CTF18-RFC (Fig S15). The CTF18-RFC mutant displayed altered concentration dependencies compared to the wild-type, thereby validating the mutant phenotype and underscoring the structural importance of the β-hairpin feature. Notably, the primer extension assay was conducted with 2 nM plasmid DNA and excess CTF18-RFC (40 nM), PCNA (100 nM), and Pol ε (30 nM), establishing single-turnover conditions per DNA molecule. Under these conditions, the rate—not the extent—of PCNA loading becomes the key determinant of synthesis dynamics. Although Pol ε can initiate synthesis independently of PCNA, the clamp is required for full processivity. Therefore, the slower loading rate observed for the CTF18-RFC mutant likely delays formation of highly processive complexes, resulting in a slower accumulation of extended products and explaining the observed ∼2-fold reduction in synthesis.

Primer extension assay with Pol ε and CTF18-RFC WT or CTF18△165-194-RFC.

a) Time course reactions on M13mp18 single-strand DNA with 40 nM wildtype CTF18-RFC or CTF18 Δ165-194-RFC. Reactions were performed as described in Materials and Methods. b) Quantification of the band intensities from the primer extension assay shows that the rate of increase in intensity, derived from the initial linear portion of the reaction time before saturation, is 4142 ± 95 a.u./min for CTF18-RFC and 2389 ± 143 a.u./min for CTF18 Δ165-194-RFC. This indicates that the mutant slows Pol ε synthesis by 42%. The experiment was conducted three times. Dots in b) represent mean values from the three replicas.

Discussion

In this work, we reported the structure of the human CTF18-RFC clamp loader bound to PCNA in the presence of ATP with or without Mg2+ ions. Our cryo-EM data, together with previous work (34) (37), support that the Ctf8 and Dcc1 subunits of CTF18-RFC, forming the regulatory module interacting with Pol ε (37), are flexibly tethered to the RFC-module via binding to the Ctf18 C-terminus, which is separated by the N-terminal RFC1-like module of Ctf18 by ∼90 amino acid flexible linker. Therefore, CTF18-RFCcomprises two structurally distinct modules that cooperate to load PCNA onto the leading strand, facilitating replication by Pol ε.

Earlier version: The CTF18-RFC−PCNA complex was observed in an autoinhibited conformation, which prevents DNA binding. This mirrors the similar conformation in hRFC and scRFC, as previously reported (10, 11). Such conformation represents a transient encounter complex that occurs early in the clamp loading process, before the opening of the clamp. Despite the variations in the large RFC subunit in the two assemblies, we demonstrate a striking architectural conservation between the alternative human clamp loader CTF18-RFCand the canonical RFC. Specifically, CTF18-RFC utilizes four subunits (Ctf18, RFC2, RFC5, and RFC4) to engage two PCNA protomers. However, the engagement of the third PCNA protomer by RFC3 is necessary to initiate the opening of the PCNA ring interface, a process that requires a conformational shift of the AAA+ pentamer from an autoinhibited to an active state. Supporting this, previous structural studies on scRFC have shown that the conformation of RFC bound to open PCNA is facilitated through a ’crab-claw’ mechanism, which is inherently predisposed to DNA binding (11). This conformational change creates an opening between the A′ segment and the AAA+ module in the A subunit, allowing primer-template DNA to bind within the RFC inner chamber (11) (Fig 4C). Remarkably, this significant conformational shift occurs rapidly (52), does not depend on ATP hydrolysis, and ensures the sequential binding of PCNA and the DNA substrate in the correct order. While our study did not resolve the active conformation of CTF18-RFC, the observed structural conservation in the autoinhibited state, particularly of the large Ctf18 subunit (Fig 2), suggests that a similar ’crab-claw’ mechanism may facilitate the transition of CTF18-RFC to its active state for PCNA opening and DNA binding. While our work was being finalized, several cryo-EM structures of human CTF18-RFC bound to PCNA, and primer/template DNA were reported by another group (53). These findings are consistent with the distinct features of CTF18-RFC observed in our structures and independently support the notion of significant mechanistic similarity between CTF18-RFC and canonical RFC in loading PCNA onto a ss/dsDNA junction.

Revised version: The CTF18-RFC–PCNA complex was observed in an autoinhibited conformation, which prevents DNA binding. This mirrors the similar conformation in hRFC and scRFC, as previously reported (10), (11). Such conformation represents a transient encounter complex that occurs early in the clamp loading process, before the opening of the clamp. Despite the variations in the large RFC subunit in the two assemblies, we demonstrate a striking architectural conservation between the alternative human clamp loader CTF18-RFCand the canonical RFC. Specifically, CTF18-RFC utilizes four subunits (Ctf18, RFC2, RFC5, and RFC4) to engage two PCNA protomers. However, the engagement of the third PCNA protomer by RFC3 is necessary to initiate the opening of the PCNA ring interface, a process that requires a conformational shift of the AAA+ pentamer from an autoinhibited to an active state. Supporting this, previous structural studies on scRFC have shown that the conformation of RFC bound to open PCNA is facilitated through a ‘crab-claw’ mechanism, which is inherently predisposed to DNA binding (11). This conformational change creates an opening between the A′ segment and the AAA+ module in the A subunit, allowing primer-template DNA to bind within the RFC inner chamber (11) (Fig 4C). Remarkably, this significant conformational shift occurs rapidly (55), does not depend on ATP hydrolysis, and ensures the sequential binding of PCNA and the DNA substrate in the correct order. While our study did not resolve the active conformation of CTF18-RFC, the observed structural conservation in the autoinhibited state, particularly of the large Ctf18 subunit (Fig 2), suggests that a similar ‘crab-claw’ mechanism may facilitate the transition of CTF18-RFC to its active state for PCNA opening and DNA binding. While our work was being finalized, several cryo-EM structures of human CTF18-RFC bound to PCNA, and primer/template DNA were reported by another group (56). These findings are consistent with the distinct features of CTF18-RFC observed in our structures and independently support the notion of significant mechanistic similarity between CTF18-RFC and canonical RFC in loading PCNA onto a ss/dsDNA junction.

Earlier version: The main structural differences between CTF18-RFC and the canonical RFC include increased mobility of the large subunit and the presence of a β-hairpin that attaches to the RFC5 subunit. The mobility of the large subunit is likely enhanced by the low-affinity PIP motif that anchors it to the PCNA ring. We propose that the β-hairpin may have evolved to counterbalance the mobility of the large subunit, helping to maintain its connection to the RFC pentamer and stabilizing the clamp loader structure. Our structural data and functional assays indeed show that deletion of the β-hairpin results in reduced stability of the CTF18-RFC−PCNA complex, slower clamp loading activity, and diminished activity of CTF18-RFC in stimulating Pol ε primer synthesis, likely due to reduced CTF18-RFC efficiency in PCNA loading (Fig 8). Interestingly, a recent work reporting cryo-EM structures of the yeast homolog of CTF18-RFC bound to PCNA (54) does not support the presence of a stable β-hairpin at the N-terminus of the Ctf18 subunit engaged to the collar of the Rfc3 subunit. Therefore, the functional β-hairpin we observed in human CTF18-RFC appears a distinct feature that pertains to higher eukaryotes, required to maximize the activity of CTF18-RFC in loading PCNA onto the leading strand.

Revised version: The main structural differences between CTF18-RFC and the canonical RFC include increased mobility of the large subunit and the presence of a β-hairpin that attaches to the RFC5 subunit. The mobility of the large subunit is likely enhanced by the low-affinity PIP motif that anchors it to the PCNA ring. We propose that the β-hairpin may have evolved to counterbalance the mobility of the large subunit, helping to maintain its connection to the RFC pentamer and stabilizing the clamp loader structure. Collectively, our structural and biochemical analyses reveal that removal of the β-hairpin compromises the stability of the CTF18-RFC–PCNA complex, results in slower clamp loading kinetics, and leads to reduced stimulation of Pol ε primer synthesis, consistent with a defect in the rate of PCNA loading (Fig 8). Interestingly, a recent work reporting cryo-EM structures of the yeast homolog of CTF18-RFC bound to PCNA (56) does not support the presence of a stable β-hairpin at the N-terminus of the Ctf18 subunit engaged to the collar of the Rfc3 subunit. Therefore, the functional β-hairpin we observed in human CTF18-RFC appears a distinct feature that pertains to higher eukaryotes, required to maximize the activity of CTF18-RFC in loading PCNA onto the leading strand.

Earlier version: Deletion of the N-terminal β-hairpin in the Ctf18 large subunit slows down PCNA loading and reduces DNA synthesis by Pol ε.

Revised version: Deletion of the N-terminal ε and CTF18-RFC WT or CTF18△165-194-RFC.β-hairpin in the Ctf18 large subunit slows down PCNA loading and reduces DNA synthesis by Pol ε.

a) The structure of the RFC module of CTF18-RFC is stabilized by the presence of a β-hairpin in the Ctf18 subunit (amino acids 165-194). Thus, wild type CTF18-RFC loads PCNA onto DNA efficiently and stimulates DNA synthesis by Pol ε interacting with the polymerase. b) The stability of the CTF18-RFC–PCNA complex is diminished by the deletion of the β-hairpin. Consequently, a CTF18-RFC deletion mutant lacking the β-hairpin (CTF18Δ165-194-RFC) shows slower loading of PCNA and lower activity in stimulating primer synthesis by Pol ε.

Materials and methods

Protein expression and purification

Preparation of CTF18-RFC

MultiBac™ expression system (Geneva Biotech) was used to express human Ctf18-RFC2-5 along with two additional cohesion-specific factors, Dcc1 and Ctf8 (named hereafter CTF18-RFC) in Sf9 insect cells. For this purpose, insect cells optimized sequences of N-terminus twin Strept-tagged TEV Ctf18, RFC2, and RFC4 were cloned in pACEBac1 plasmid with independent promoters and terminators, and N-terminus 6X histidine-tagged RFC3, RFC5, Dcc1 and Ctf8 were cloned in pIDS plasmid with independent promoters and terminators by GenScript. Finally, the single transfer vector with different subunit assemblies was generated using cre recombinase according to the MultiBac™ expression system user manual. CTF18-RFC encoding all seven subunits in a single MultiBac™ expression plasmid (pACEBac1) was transformed into DH10MultiBac™ cells to make bacmid DNA. To prepare the baculovirus, bacmid DNA containing all subunits was transfected into Sf9 cells using FuGENE® HD (Promega) according to the manufacturer’s instructions. This baculovirus prep was amplified twice to obtain a higher titer virus (P3 virus). The expression of CTF18-RFC then proceeded by transfecting a 6 L Hi5 suspension culture at a density of 2 × 106 cells/mL with the P3 virus for 60-65 hours.

Cells were collected by centrifugation at 5,500 xg for 10 minutes and re-suspended in lysis buffer [50 mM Tris-HCl (pH 7.5), 300 mM NaCl, 20 mM imidazole, 5 mM BME, 0.1% Nonidet P-40, 10% glycerol and one EDTA-free protease inhibitor cocktail tablet/50 mL]. All further steps were performed at 4° C. Cells were sonicated and debris was removed by centrifugation at 95,834 xg for 1 hour at 4 °C. The resultant supernatant was directly loaded onto a HisTrap HP 5 mL affinity column (Cytiva) pre-equilibrated with buffer A [50 mM Tris-HCl (pH 7.5), 300 mM NaCl, 20 mM imidazole, 5 mM BME and 10% glycerol]. After loading, the column was washed with 50 mL of buffer A and the bound fractions were eluted by gradient with 50 mL of buffer B [50 mM Tris-HCl (pH 7.5), 300 mM NaCl, 500 mM imidazole, 5 mM BME, and 10% glycerol]. The eluents were pooled and directly loaded onto Strept Trap XT 1 mL affinity column pre-equilibrated with buffer A [50 mM Tris-HCl (pH 7.5), 300 mM NaCl, 5 mM BME and 10% glycerol]. After loading, the column was washed with 10 mL of buffer A, followed by step elution with buffer B [50 mM Tris-HCl (pH 7.5), 300 mM NaCl, 50 mM Biotin, 5 mM BME, and 10% glycerol]. Eluents were pooled and incubated with the TEV protease overnight. Fractions were concentrated to 1 mL and loaded onto HiLoad 16/600 Superdex 200 pg pre-equilibrated with gel filtration buffer [50 mM Tris-HCl (pH 7.5), 200 mM NaCl, 1 mM DTT, and 10% glycerol]. Protein fractions were pooled in a total volume of 1.6 mL, flash frozen, and stored at −80°C.

Preparation of PCNA

Earlier version: To express PCNA BL21(DE3) E. coli cells were transformed with WT PCNA plasmid. The transformed cells were grown at 37°C in 2YT media supplemented with ampicillin to an OD600 of 1.2. The cell cultures were induced with 0.5 mM isopropyl-beta-D-1-thiogalactopyranoside (IPTG) and let them grow for 19 hours at 16°C. Cells were then harvested by centrifugation at 5500 xg for 10 minutes and then re-suspended in lysis buffer [50 mM Tris-HCl (pH 7.5), 750 mM NaCl, 20 mM imidazole, 5 mM BME, 0.2% Nonidet P-40, 1 mM PMSF, 5% glycerol and one EDTA-free protease inhibitor cocktail tablet/50 mL]. The cells were lysed with 2 mg/mL of lysozyme at 4°C for 30 minutes. Cells were sonicated and debris was removed by centrifugation at 22,040 xg for 20 minutes at 4°C. The cleared cell lysate was directly loaded onto a HisTrap HP 5 mL affinity column (Cytiva) pre-equilibrated with buffer A [50 mM Tris-HCl (pH 7.5), 500 mM NaCl, 20 mM imidazole, 5 mM BME and 5% glycerol]. The column was firstly washed with 50 mL of buffer A and then with 50 mL of buffer B containing low salt [50 mM Tris-HCl (pH 7.5), 100 mM NaCl, 500 mM imidazole, 5 mM BME, and 5% glycerol]. The eluents were pooled and directly loaded onto HiTrap Q HP 5 mL anion exchange column (GE Healthcare) pre-equilibrated with buffer D [50 mM Tris-HCl (pH 7.5), 100 mM NaCl, 5 mM BME and 5% glycerol]. The column was washed with 50 mL of buffer D after loading, followed by elution of 50 mL gradient with buffer E [50 mM Tris-HCl (pH 7.5), 1 M NaCl, 5 mM BME, and 5% glycerol]. Eluents were pooled and concentrated to 1.5 mL and loaded onto HiLoad 16/600 Superdex 200 pg pre-equilibrated with gel filtration buffer [50 mM Tris-HCl (pH 7.5), 150 mM NaCl, 1 mM DTT, and 5% glycerol]. Protein fractions were pooled, flash-frozen, and stored at -80 °C. Human RPA, Cy5-PCNA, and RFC for the clamp loading assay were obtained as previously described (55, 56). The concentration of active RPA was determined via a FRET-based activity assay as described previously (57).

Revised version: To express PCNA BL21(DE3) E. coli cells were transformed with WT PCNA plasmid. The transformed cells were grown at 37°C in 2YT media supplemented with ampicillin to an OD600 of 1.2. The cell cultures were induced with 0.5 mM isopropyl-beta-D-1-thiogalactopyranoside (IPTG) and let them grow for 19 hours at 16°C. Cells were then harvested by centrifugation at 5500 xg for 10 minutes and then re-suspended in lysis buffer [50 mM Tris-HCl (pH 7.5), 750 mM NaCl, 20 mM imidazole, 5 mM BME, 0.2% Nonidet P-40, 1 mM PMSF, 5% glycerol and one EDTA-free protease inhibitor cocktail tablet/50 mL]. The cells were lysed with 2 mg/mL of lysozyme at 4°C for 30 minutes. Cells were sonicated and debris was removed by centrifugation at 22,040 xg for 20 minutes at 4°C. The cleared cell lysate was directly loaded onto a HisTrap HP 5 mL affinity column (Cytiva) pre-equilibrated with buffer A [50 mM Tris-HCl (pH 7.5), 500 mM NaCl, 20 mM imidazole, 5 mM BME and 5% glycerol]. The column was firstly washed with 50 mL of buffer A and then with 50 mL of buffer B containing low salt [50 mM Tris-HCl (pH 7.5), 100 mM NaCl, 500 mM imidazole, 5 mM BME, and 5% glycerol]. The eluents were pooled and directly loaded onto HiTrap Q HP 5 mL anion exchange column (GE Healthcare) pre-equilibrated with buffer D [50 mM Tris-HCl (pH 7.5), 100 mM NaCl, 5 mM BME and 5% glycerol]. The column was washed with 50 mL of buffer D after loading, followed by elution of 50 mL gradient with buffer E [50 mM Tris-HCl (pH 7.5), 1 M NaCl, 5 mM BME, and 5% glycerol]. Eluents were pooled and concentrated to 1.5 mL and loaded onto HiLoad 16/600 Superdex 200 pg pre-equilibrated with gel filtration buffer [50 mM Tris-HCl (pH 7.5), 150 mM NaCl, 1 mM DTT, and 5% glycerol]. Protein fractions were pooled, flash-frozen, and stored at −80 °C. Human RPA, Cy5-PCNA, and RFC for the clamp loading assay were obtained as previously described (57) (58). The concentration of active RPA was determined via a FRET-based activity assay as described previously (59).

Preparation of DNA polymerase ε

MultiBac™ expression system (Geneva Biotech) was used to express human DNA Pol ε in Sf9 insect cells. For this purpose, insect cells optimized sequence of the catalytic subunit, p261 was cloned in pACEBac1 and p59, C-terminus TEV twin Strept-tagged p17, and N-terminus 8X histidine-tagged TEV p12 along with independent promoters and terminators were cloned in pIDS by the GenScript. Finally, the single transfer vector with different subunit assemblies was generated using cre recombinase according to MultiBac™ expression system user manual. To prepare the baculovirus, the bacmid DNA of Pol ε was transfected into Sf9 cells using FuGENE® HD (Promega) according to the manufacturer’s instructions. The resulting supernatant was obtained as the P1 virus stock which was then amplified to obtain P2 virus stock. P2 virus stock was then further amplified to obtain P3 virus stock for large-scale expression. Pol ε was expressed by transfecting 4 L of Sf9 suspension culture at 2 × 106 cells/mL density with P3 virus. Cells were harvested after 72 hours by centrifugation at 5,500 xg for 10 minutes and re-suspended in 200 mL of lysis buffer [50 mM HEPES (pH 8), 20 mM Imidazole, 400 mM NaCl, 5 mM β-Mercaptoethanol, 5% (v/v) Glycerol, 0.1% NP-40, 1 mM PMSF and EDTA-free protease inhibitor cocktail tablet/50 mL (Roche, UK)]. Cells were sonicated and debris was removed by centrifugation at 95,834 xg for 1 hour at 4 °C. The resultant supernatant was directly loaded onto a HisTrap HP 5 mL affinity column (Cytiva) equilibrated with buffer A [50 mM HEPES (pH 8), 20 mM Imidazole, 400 mM NaCl, 5mM β-Mercaptoethanol, 5% (v/v) Glycerol and EDTA-free protease inhibitor cocktail tablet/50 mL (Roche, UK]. After loading, the column was washed with 50 mL of buffer A and the bound fractions were eluted by gradient with 50 mL of buffer B [50 mM HEPES (pH 8), 500 mM Imidazole, 300 mM NaCl, 5 mM β-Mercaptoethanol, 5% (v/v) Glycerol and EDTA-free protease inhibitor cocktail tablet/50 ml (Roche, UK)]. Peak fractions containing Pol ε subunits were pooled and loaded directly onto Strep Trap XT 5 mL (Cytiva) column equilibrated with buffer C [50 mM HEPES (pH 8), 250 mM NaCl, 5 mM β-Mercaptoethanol, 5% (v/v) Glycerol] and EDTA-free protease inhibitor cocktail tablet/50 mL (Roche, UK)]. After loading, the column was washed with 50 mL of buffer C and the bound fractions were incubated for 2-4 hours with 15 mL of 100% buffer D [50 mM HEPES (pH 8), 100 mM NaCl, 50 mM Biotin, 5 mM β-Mercaptoethanol, 5% (v/v) Glycerol] and EDTA-free protease inhibitor cocktail tablet/50 mL (Roche, UK)] and eluted further with 30 mL of 100% buffer D. Fractions that contain all Pol ε subunits were pooled and incubated with the TEV protease for overnight and loaded onto an anion exchanger, Mono Q 5/50 GL column (Cytiva) pre-equilibrated with buffer E [50 mM HEPES (pH 8), 100 mM NaCl, 5 mM β-Mercaptoethanol and 5% Glycerol]. After loading, the column was washed with 15 mL of buffer E. The Pol ε was eluted with a 10 mL gradient from 100 mM NaCl to 1 M NaCl in 50 mM HEPES (pH 8), 5 mM β-Mercaptoethanol, and 5% Glycerol. The fractions that contained all Pol ε subunits were combined and concentrated to 1 mL and loaded onto HiLoad 16/600 Superdex 200 pg (Cytiva) pre-equilibrated with gel filtration buffer [50mM HEPES (8), 200 mM NaCl, 1 mM DTT and 5% Glycerol]. Protein fractions were pooled, flash-frozen, and stored at −80°C.

All protein concentrations were calculated by measuring their absorbance at 280 nm using the extinction coefficients calculated from the amino acid composition (CTF18-RFC: 216,360 M−1 cm −1; PCNA: 47,790 M−1 cm −1, Pol ε: 351,470 M−1 cm −1).

DNA substrates for Cryo-EM

For the primer/template substrates used in all the structures described here, 80 base pairs template strand 5’-TTTTTTTTTTTTTTTTTTTTTTTTTTTTTTTTTTTTTTTTGACTGCAC GAATTAAGCAATTCGTAATCATGG TCATAGCT-3’ was annealed to a 39 base pairs primer containing a 3’ dideoxy cytosine chain terminator 5’-AGCTATG ACCATGATTACGAATTGCTTAATTCGTGCAGT[ddC]-3’ to form the P/T substrate. The strands were mixed in an equimolar ratio in the presence of 20 mM Tris pH 7.5 and 25 mM NaCl. The annealing reaction was performed by incubating the P/T at 92°C for 2 minutes followed by slow cooling to room temperature overnight. All oligonucleotides were purchased from IDT Integrated DNA Technologies.

Clamp loading and primer extension substrates

Earlier version: Oligonucleotides for the clamp loading assay comprising the 5′ddPCy3/T P/T DNA substrate were synthesized by Integrated DNA Technologies (Coralville, IA) and purified on denaturing polyacrylamide gels. The concentration of unlabeled template strand DNA was determined from the absorbance at 260 nm using the provided extinction coefficient. Concentrations of the Cy3-labeled primer strand DNA was determined from the extinction coefficient of Cy3 at 550 nm (ε550 = 136,000 M−1cm−1). For annealing the 5′ddPCy3/T P/T DNA substrate, the Cy3-labled primer strand DNA and the unlabeled template strand DNA were mixed in equimolar amounts in 1X annealing buffer (10 mM TrisHCl, pH 8.0, 100 mM NaCl, 1 mM EDTA), heated to 95 °C for 5 minutes, and allowed to slowly cool to room temperature. For the prime extension assay, Oligo (592:TAACGCCAGGGTTTTCCCAGTCACG) (Integrated DNA Technologies) was annealed to 50 nM M13mp18 single-stranded DNA (New England Biolabs) in a reaction buffer consisting of [10 mM Tris-HCl (pH 7.6), 100 mM NaCl and 5 mM EDTA]. The mixture was heated to 95 °C for 5 minutes, then cooled gradually to room temperature. The unannealed oligonucleotide was removed using a QIAquick PCR purification kit (Qiagen).

Revised version: Oligonucleotides for the clamp loading assay comprising the 5′ddPCy3/T P/T DNA substrate were synthesized by Integrated DNA Technologies (Coralville, IA) or Bio-Synthesis (Lewisville, TX) and purified on denaturing polyacrylamide gels. The concentrations of unlabeled DNAs were determined from the absorbance at 260 nm using the provided extinction coefficients. Concentrations of Cy3-labeled DNAs were determined from the extinction coefficient of Cy3 at 550 nm (ε550 = 136,000 M−1cm−1). For annealing the 5′ddPCy3/T P/T DNA substrate (as depicted in Fig S10), the primer and corresponding complementary template strand were mixed in equimolar amounts in 1X annealing buffer (10 mM TrisHCl, pH 8.0, 100 mM NaCl, 1 mM EDTA), heated to 95 °C for 5 minutes, and allowed to slowly cool to room temperature. For the prime extension assay, Oligo (592:TAACGCCAGGGTTTTCCCAGTCACG) (Integrated DNA Technologies) was annealed to 50 nM M13mp18 single-stranded DNA (New England Biolabs) in a reaction buffer consisting of [10 mM Tris-HCl (pH 7.6), 100 mM NaCl and 5 mM EDTA]. The mixture was heated to 95 °C for 5 minutes, then cooled gradually to room temperature. The unannealed oligonucleotide was removed using a QIAquick PCR purification kit (Qiagen).

Pre-steady state FRET measurements

Earlier version: All experiments were performed at room temperature (23 ± 2 °C) in a 16.100F-Q-10/Z15 sub-micro fluorometer cell (Starna Cells) and monitored in a Horiba Scientific Duetta-Bio fluorescence/absorbance spectrometer as described previously (57). In short, reaction solutions are excited at 514 nm and the fluorescence emission intensities (I) are monitored essentially simultaneously (Δt = 0.118 ms) at the peak emission wavelengths for Cy3 (563 nm, I563) and Cy5 (665 nm, I665) over time, recording I every 0.17 s. For all experiments, excitation and emission slit widths are 10 nm. All recorded fluorescence emission intensities are corrected by a respective dilution factor and all-time courses are adjusted for the time between the addition of each component and the fluorescence emission intensity recordings (i.e., dead time = 6 s). For any recording of the fluorescence emission intensities (I665 and I563), the approximate FRET efficiency is estimated from the equation For each experiment below, the final concentrations of all reaction components are indicated.

Revised version: All experiments were performed at room temperature (23 ± 2 °C) in a 16.100F-Q-10/Z15 sub-micro fluorometer cell (Starna Cells) and monitored in a Horiba Scientific Duetta-Bio fluorescence/absorbance spectrometer as described previously (52). In short, reaction solutions are excited at 514 nm and the fluorescence emission intensities (I) are monitored essentially simultaneously (Δt = 0.118 ms) at the peak emission wavelengths for Cy3 (563 nm, I563) and Cy5 (665 nm, I665) over time, recording I every 0.17 s. For all experiments, excitation and emission slit widths are 10 nm. All recorded fluorescence emission intensities are corrected by a respective dilution factor and all-time courses are adjusted for the time between the addition of each component and the fluorescence emission intensity recordings (i.e., dead time ≤ 7 s). For any recording of the fluorescence emission intensities (I665 and I563), the approximate FRET efficiency is estimated from the equation . For each experiment below, the final concentrations of all reaction components are indicated.

Earlier version: PCNA loading experiments were carried out in 1X Replication Buffer (25 mM HEPES, pH 7.5, 125 mM KOAc, 10 mM Mg(OAc)2) supplemented with 1 mM DTT, 1 mM ATP, and the ionic strength adjusted to physiological (200 mM) by addition of KOAc. Experiments were performed via slight modifications to a published protocol (58, 59). In short, 5′ddPCy3/T P/T DNA (200 nM, NeutrAvidin (800 nM homotetramer), and ATP (1 mM) are pre-incubated with RPA (600 nM heterotrimer). Then, Cy5-PCNA (200 nM homotrimer) is added, the resultant solution is transferred to a fluorometer cell, and the cell is placed in the instrument. I665 and I563 are monitored until both stabilize for at least 1 min. The EFRET values within this stable region represent the signal for Cy5-PCNA remaining completely disengaged from 5¢ddPCy3/T•RPA nucleoprotein complex and are averaged to obtain the EFRET value observed prior to addition of a clamp loader complex. Finally, a clamp loader complex (200 nM RFC, CTF18, or CTF18Δ165-194 -RFC heteropentamer) is added, the resultant solution is mixed, and I665 and I563 are monitored, beginning 6 s after the addition of the respective clamp loader complex (i.e., Δt = 6 s).

Revised version: PCNA loading experiments were carried out in 1X Replication Buffer (25 mM HEPES, pH 7.5, 125 mM KOAc, 10 mM Mg(OAc)2) supplemented with 1 mM DTT, 1 mM ATP, and the ionic strength adjusted to physiological (200 mM) by addition of KOAc. Experiments were performed via slight modifications to a published protocol (52) (53). In short, 5′ddPCy3/T P/T DNA (200 nM, NeutrAvidin (800 nM homotetramer), and ATP (1 mM) are pre-incubated with RPA (600 nM heterotrimer). Then, Cy5-PCNA (200 Nm homotrimer) is added, the resultant solution is transferred to a fluorometer cell, and the cell is placed in the instrument. I665 and I563are monitored until both stabilize for at least 1 min. Finally, a given clamp loader complex (200 nM RFC, CTF18-RFC, or CTF18D165-194-RFC henteropentamer) is added, the resultant solution is mixed, and I665 and I563are monitored, beginning ≤ 7 s after the addition of the respective clamp loader complex (i.e., “dead time” = Dt ≤ 7 s).

Primer extension assay

The primer extension reaction was conducted in a reaction buffer containing [40 mM HEPES (pH 7.6), 100 mM K-glutamate, 1 mM DTT, 10 mM Mg-Acetate, 200 μg/mL BSA, 1 mM ATP, 100 μM each of dATP, dTTP, and dGTP, 20 μM dCTP, and 10 μCi of [32P]-dCTP (Hartmann)]. Two nanomolar of the primed-M13 template were initially incubated at 37 °C for 10 minutes with 100 nM PCNA, 400 nM RPA, and either CTF18-RFCWT or CTF18Δ165-194-RFC. The reaction was initiated by the addition of 30 nM Pol ε. Reactions were stopped at the indicated time point by adding 50 mM EDTA and electrophoresed on 0.7% alkaline agarose gel at 35 V for 17 hours. The gel was backed with DE81 paper, compressed, and imaged in a Sapphire Biomolecular Imager (Azure Biosystems). All the experiments were performed in triplicate.

Cryo-EM grids preparation and data collection

For all the complexes (WT and mutant), UltraAuFoil® R1.2/1.3 Au 300 grids were glow-discharged at 30 mA for 30 seconds on an EasyGlow glow-discharge unit. For the preparation of the complex without Mg2+ (Dataset 1 and Dataset 3), P/T DNA, ATP, CTF18-RFC WT/CTF18Δ165-194-RFC, and PCNA were mixed in order. The final concentrations were 5.6 μM P/T DNA, 0.2 mM ATP, 1.5 μM CTF18-RFC WT/CTF18Δ165-194-RFC, and 7.5 μM PCNA. The buffer this was performed in comprised 50 mM Tris-HCl pH 7.5, 200 mM NaCl, 1 mM DTT, and 0.5 mM ATP. 3 μL of the sample was applied to the grid, blotted for 2 s at blot force −5, and plunged frozen into liquid ethane using a Vitrobot Mark IV (FEI Thermo Fisher), which was set at 4°C and 100% humidity. For the preparation of CTF18-RFC/PCNA/P-T DNA/ATP/Mg2+complex (Dataset 2), the components P/T DNA, ATP, Mg2+, CTF18-RFC, and PCNA were mixed in order. The final concentrations were 5.6 μM P/T DNA, 0.2 mM ATP, 1.9 mg/mL Mg2+, 1.5 μM CTF18-RFC, and 7.5 μM PCNA. The buffer this was performed in comprised 50 mM Tris-HCl pH 7.5, 200 mM NaCl, 1 mM DTT, 0.5 mM ATP, and 5 mM Mg2+. 3 μL of the sample was applied to the grid in the same way as above for data collection. Cryo-EM data for all the samples were collected on a Thermo Fisher Scientific Titan Krios G4 transmission electron microscope at KAUST. Electron micrographs for Dataset 1 (CTF18-RFC-PCNA with ATP in the absence of Mg2+), were collected using a Falcon 4i detector at a dose rate of 7.5 e−/pix/sec for 5 seconds and a calibrated pixel size of 0.93 Å. Focusing was performed over a range between −2.7 μm and −1.5μm, in 0.3 μm intervals. Electron micrographs for Dataset 2 (CTF18-RFC-PCNA with ATP and Mg2+) were collected using a Falcon 4i detector at a dose rate of 8.02 e−/pix/sec for 5 seconds and a calibrated pixel size of 0.93 Å. Focusing was performed over a range between - 2.7 μm and −1.5 μm, in 0.3 μm intervals. Electron micrographs for Dataset 3 (CTF18Δ165-194-RFC-PCNA with ATP and in the absence of Mg2+) were collected using a Falcon 4i detector at a dose rate of 7.4 e−/pix/sec for 5 seconds and a calibrated pixel size of 0.93 Å. Focusing was performed over a range between −2.7 μm and −1.5 μm, in 0.3 μm intervals.

Cryo-EM image processing

Earlier version: Preprocessing of all datasets (Dataset 1: CTF18-RFC, ATP, primer/template DNA and PCNA; Dataset 2: CTF18-RFC, ATP, Mg2+, primer/template DNA and PCNA; Dataset 3: CTF18Δ165-194-RFC, ATP, primer/template DNA and PCNA) was performed as follows: the micrographs were corrected for beam-induced motion and then integrated using MotionCor2 (60). All frames were retained and a patch alignment of 4x4 was used. CTFFIND-4.12 estimated Contrast Transfer Function (CTF) parameters for each micrograph (61). Integrated movies were inspected with Relion-4.0 (Dataset 1 and Dataset 2) and Relion-5.0 (Dataset 3) for further image processing (7906 movies for Dataset 1, 6181 movies for Dataset 2, and 6499 movies for Dataset 3) (62). Particle picking was performed in automated mode using Topaz (63) for Dataset 1. Particle extraction was carried out from micrographs using a box size of 55 pixels (pixel size: 3.65 Å/pixel). An initial dataset of 2 x 106 particles was cleaned by 2D classification followed by 3D classification with alignment. Three 3D classes were generated with populations of 31, 28, and 40%. 3D refinement, CtfRefine, and polishing were performed on the selected 3D classes corresponding to populations 31 and 40%. A bigger box size, 480 pixels (pixel size: 0.93 Å/pixel) was used to re-extract the particles. 3D refinement, CtfRefine, and polishing yielded reconstruction at 2.9 Å.

Revised version: Preprocessing of all datasets (Dataset 1: CTF18-RFC, ATP, primer/template DNA and PCNA; Dataset 2: CTF18-RFC, ATP, Mg2+, primer/template DNA and PCNA; Dataset 3: CTF18Δ165-194-RFC, ATP, primer/template DNA and PCNA) was performed as follows: the micrographs were corrected for beam-induced motion and then integrated using MotionCor2 (60). All frames were retained and a patch alignment of 4×4 was used. CTFFIND-4.12 estimated Contrast Transfer Function (CTF) parameters for each micrograph (61). Integrated movies were inspected with Relion-4.0 (Dataset 1 and Dataset 2) and Relion-5.0 (Dataset 3) for further image processing (7906 movies for Dataset 1, 6181 movies for Dataset 2, and 6499 movies for Dataset 3) (62). Particle picking was performed in automated mode using Topaz (63) for Dataset 1. Particle extraction was carried out from micrographs using a box size of 55 pixels (pixel size: 3.65 Å/pixel). An initial dataset of 2 × 106 particles was cleaned by 2D classification followed by 3D classification with alignment. Three 3D classes were generated with populations of 31, 28, and 40%. 3D refinement, CtfRefine, and polishing were performed on the selected 3D classes corresponding to populations 31 and 40%. A bigger box size, 480 pixels (pixel size: 0.93 Å/pixel) was used to re-extract the particles. 3D refinement, CtfRefine, and polishing yielded reconstruction at 2.9 Å.

Earlier version: Particle picking was performed in reference-based mode for Dataset 2. Particle extraction was carried out from micrographs using a box size of 55 pixels (pixel size of 3.65 Å/pixel). An initial dataset of 3 x 106 particles was cleaned by 2D classification followed by 3D classification with alignment. Three 3D classes were generated with 30, 33, and 36% populations. 3D refinement, CtfRefine, and polishing were performed on the selected 3D classes corresponding to populations 33 and 36%. 3D refinement, CtfRefine, and polishing yielded reconstruction at 3.2 Å. Particle picking was performed in automated mode using Topaz (63) for Dataset 3. Particle extraction was done from micrographs using a box size of 220 pixels (pixel size: 0.93 Å/pixel). An initial dataset of 1 x 106 particles was cleaned by 2D classification. For quantification of complex stability, the number of particles contributing to each 2D class was extracted from the classification metadata (Datasets 1 and 3). All classes showing isolated PCNA rings were summed and compared to the total number of particles in classes representing intact CTF18-RFC– PCNA complexes. This analysis was performed for both wild-type and β-hairpin deletion mutant datasets. Notably, no 2D classes corresponding to free PCNA were observed in the wild-type dataset, whereas in the mutant dataset, a substantial fraction of particles corresponded to isolated PCNA, suggesting reduced stability of the mutant complex.

Revised version: Particle picking was performed in reference-based mode for Dataset 2. Particle extraction was carried out from micrographs using a box size of 55 pixels (pixel size of 3.65 Å/pixel). An initial dataset of 3 × 106 particles was cleaned by 2D classification followed by 3D classification with alignment. Three 3D classes were generated with 30, 33, and 36% populations. 3D refinement, CtfRefine, and polishing were performed on the selected 3D classes corresponding to populations 33 and 36%. 3D refinement, CtfRefine, and polishing yielded reconstruction at 3.2 Å. Particle picking was performed in automated mode using Topaz (63) for Dataset 3. Particle extraction was done from micrographs using a box size of 220 pixels (pixel size: 0.93 Å/pixel). An initial dataset of 1 × 106 particles was cleaned by 2D classification. For quantification of complex stability, the number of particles contributing to each 2D class was extracted from the classification metadata (Datasets 1 and 3). All classes showing isolated PCNA rings were summed and compared to the total number of particles in classes representing intact CTF18-RFC–PCNA complexes. This analysis was performed for both wild-type and β-hairpin deletion mutant datasets. Notably, no 2D classes corresponding to free PCNA were observed in the wild-type dataset, whereas in the mutant dataset, a substantial fraction of particles corresponded to isolated PCNA, suggesting reduced stability of the mutant complex.

Cryo-EM 3D Variability Analysis

3DVA (43) (44) of the CTF18-RFC–PCNA consensus maps, with and without Mg²⁺, was performed in CryoSPARC v 4.6 using their respective refinement masks, employing three modes and a resolution filter of 5 Å, based on global FSC resolutions of 2.9 Å and 3.2 Å, respectively. For each mode, 3 frames were generated and morphed to visualize the conformational flexibility across the components.

Molecular Modelling

The structure of the Ctf18 large subunit of human CTF18-RFC (Q8WVB6, residues 277-805) was generated with Phyre2 (64), using as a template the sequence of human RFC1 (P35251, residues 581-1034). The Ctf18 subunit model was then aligned to the RFC1 subunit in the structure of the human RFC clamp loader bound to PCNA in an autoinhibited conformation (PDB 6VVO, (10)) using Chimera, the RFC1 subunit was deleted, and the resulting model including Ctf18, subunits RFC2-5 and PCNA was rigid-body fitted into the cryo-EM map. The model was edited and real-space refined with Coot (65) following flexible fitting performed with ISOLDE (66). A final real-space refinement was performed in Phenix (67), applying secondary structure constraints. The models of both CTF18-RFC–PCNA complexes with and without the Mg2+ ion were generated similarly.